# Module 1 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).
3. Do not change the title (i.e. file name) of this notebook.
4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).
5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.

-----

```python
import pandas as pd
import numpy as np

from nose.tools import assert_equal, assert_almost_equal, assert_true, assert_is_instance
```

-----

## Predicting the Price of a Car

In this assignment, we will use a partially cleaned dataset to make a predictive model. Before we attempt to build a mode, we first must load the data, select the independent feature, and the dependent label. The first Code cell below reads the data from a CSV file and displays several random instances. The second Code cell selects the independent features (`make`) and the dependent label (`price`), and displays the first few instances in this new DataFrame. 

-----

```python
df = pd.read_csv('./imports-85.data')
df.sample(5)
```

Output:
```
symboling  normalized-losses  make  fuel-type aspiration  num-of-doors  \
32           1                101     5          1        std           2.0   
129          1                122    15          1        std           2.0   
168          2                134    19          1        std           2.0   
189          3                122    20          1        std           2.0   
53           1                113     8          1        std           4.0   

     body-style drive-wheels 'engine-location'  wheel-base  ...  engine-size  \
32            3          fwd             front        93.7  ...           79   
129           3          rwd             front        98.4  ...          203   
168           0          rwd             front        98.4  ...          146   
189           4          fwd             front        94.5  ...          109   
53            2          fwd             front        93.1  ...           91   

     fuel-system  bore  stroke compression-ratio'
```

```python
df_simple = df[['make', 'price']]
df_simple.head()
```

Output:
```
make    price
0     0  13495.0
1     0  16500.0
2     0  16500.0
3     1  13950.0
4     1  17450.0
```

-----

## Problem 1: Creating the Training and Testing Datasets

The `data_split` function shown below accepts a parameter called `data` that is stored in a DataFrame. Your task is to use the `train_test_split` function available in the scikit learn library to split this DataFrame into two separate DataFrames (a testing set and training set). 

To complete this process, do the following:
- Split the training and testing set into two separate DataFrames.
- The `test_size` argument in `train_test_split` should be set to the `size` parameter.
- The `random_state` argument in `train_test_split` should be set the `random_state` parameter.
- Return both the training set and the testing set, in this order.

-----

<!-- grade_id: p1_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

def data_split(data, size=0.25, random_state=0):
    '''    
    Split data into training and testing sets.
    
    Parameters
    ---------
    data: Pandas DataFrame
    size: ratio of training and testing data
    random_state: random seed for random number generator

    Returns
    -------
    Two DataFrames: train and test, in that order.
    '''
    
    ### YOUR CODE HERE
    
    train_set, test_set = train_test_split(data, test_size=size, random_state=random_state)
    
    return train_set, test_set
```

```python
data_split(df_simple)
```

Output:
```
(     make    price
 163    19   8058.0
 61      8  10595.0
 75     10  16503.0
 106    12  18399.0
 63      8  10795.0
 ..    ...      ...
 67      9  25552.0
 192    20  13845.0
 117    13  18150.0
 47      7  32250.0
 172    19  17669.0
 
 [153 rows x 2 columns],
      make         price
 52      8   6795.000000
 181    19  15750.000000
 5       1  15250.000000
 18      3   5151.000000
 188    20   9995.000000
 170    19  11199.000000
 76     11   5389.000000
 154    19   7898.000000
 104    12  17199.000000
 33      5   6529.000000
 12      2  20970.000000
 129    15  13207.129353
 55      8  10945.000000
 66      8  18344.000000
 45      6  13207.129353
 169    19   9989.000000
 130    16   9295.000000
 7       1  18920.000000
 37      5   7895.000000
 152    19   6488.000000
 80     11   9959.000000
 111    13  15580.000000
 131    16   9895.000000
 171    19  11549.000000
 179    19  15998.000000
 138    18   5118.000000
 156    19   6938.000000
 113    13  16695.000000
 161    
```

```python
#from sklearn.model_selection import train_test_split
```

```python
# X = df_simple.iloc[:,0]
# y = df_simple.iloc[:,1]
```

```python
#X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25, random_state=0)
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
# Call function
train, test = data_split(df_simple)

# Test Data types
assert_equal(type(train), pd.DataFrame, msg="train is not a DataFrame")
assert_equal(type(test), pd.DataFrame, msg="test is not a DataFrame")

# Converting to NumPy arrays for sklearn
train_X = train.make.values.reshape(len(train.make),1)
train_y = train.price.values
test_X = test.make.values.reshape(len(test.make), 1)
test_y = test.price.values

# Test dependent values
assert_almost_equal(np.sum(test_y), 689311.38805, places=2)
assert_almost_equal(np.sum(train_y), 2018150.12935, places=2)

# Test independent values
assert_equal(train_X[1][0], 8)
assert_equal(train_X[45][0], 13)
assert_equal(test_X[25][0], 18)
assert_equal(test_X[51][0], 12)
```

-----

## Problem 2: Performing Linear Regression

Your task for this problem is to build and use the scikit learn library's `LinearRegression` estimator to  make predictions on the cars dataset. The framework for a regression function has been provided below, that takes two NumPy arrays containing the features (`trainX`) and the labels (`trainY`), an optional Boolean flag, `fit_intercept`, which indicates whether an intercept term should be fit as part of the linear regression. To complete this function, you must explicitly:
- Create a `LinearRegression` estimator by using scikit learn.
- Fit the `LinearRegression` estimator using trainX and trainY.
- Return the resulting estimator.

-----

```python
train_X[:5]
```

Output:
```
array([[19],
       [ 8],
       [10],
       [12],
       [ 8]], dtype=int64)
```

```python
train_y[:5]
```

Output:
```
array([ 8058., 10595., 16503., 18399., 10795.])
```

```python
from sklearn.linear_model import LinearRegression
lr = LinearRegression(fit_intercept=False)
```

```python
model = lr.fit(train_X, train_y)

model
```

Output:
```
LinearRegression(fit_intercept=False)
```

```python
model.get_params()
```

Output:
```
{'copy_X': True,
 'fit_intercept': False,
 'n_jobs': None,
 'normalize': 'deprecated',
 'positive': False}
```

<!-- grade_id: p2_answer, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def regression(trainX, trainY, fit_intercept=False):
    '''
    Compute a linear regression model for given training data set.
    
    Parameters
    ---------
    trainX: the training indepedent features
    trainY: the training depedent features
    fit_intercept: optional Boolean flag to indicate if an intercept should be fit

    Returns
    -------
    The fitted linear regression model
    '''
    
    ### YOUR CODE HERE
    
    lr = LinearRegression(fit_intercept=fit_intercept, n_jobs=1, normalize=False, copy_X=True, positive=False)
    lr.fit(trainX, trainY)
    
    return lr
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
lr = regression(train_X, train_y, fit_intercept=True)

assert_equal(type(lr), type(LinearRegression()))
assert_equal(lr.get_params(), {'copy_X': True, 'fit_intercept': True, \
                               'n_jobs': 1, 'normalize': False})
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\_base.py:148: FutureWarning: 'normalize' was deprecated in version 1.0 and will be removed in 1.2. Please leave the normalize parameter to its default value to silence this warning. The default behavior of this estimator is to not do any normalization. If normalization is needed please use sklearn.preprocessing.StandardScaler instead.
  warnings.warn(
File C:\ProgramData\Anaconda3\lib\unittest\case.py:1136, in TestCase.assertDictEqual(self, d1, d2, msg)
   1132 diff = ('\n' + '\n'.join(difflib.ndiff(
   1133                pprint.pformat(d1).splitlines(),
   1134                pprint.pformat(d2).splitlines())))
   1135 standardMsg = self._truncateMessage(standardMsg, diff)
-> 1136 self.fail(self._formatMessage(msg, standardMsg))

File C:\ProgramData\Anaconda3\lib\unittest\case.py:676, in TestCase.fail(self, msg)
    674 def fail(self, msg=None):
    675     """Fail immediately, with the given message."""
--> 676     raise
```

-----

## Problem 3: Checking R2 Score on Testing Dataset

For this problem, you will compute the R2 score given a model, the independent features (`X`), and the dependent feature (`y`). By using the function template provided below, you must explicitly:
- Compute the R2 score for the supplied model.
- Return the resulting score.

-----

<!-- grade_id: p3_answer, points: 0, grade: false, solution: true -->
```python
def r2_score(model, X, y):
    '''
    Compute the R2 score for a given model and data set.
    
    Parameters
    ---------
    model: linear regression model
    X: NumPy array containing indepenent data (features)
    y: NumPy array containing depenent data (labels)

    Returns
    -------
    A float containing the model score
    '''

    ### YOUR CODE HERE
    
    from sklearn.metrics import r2_score
    
    # Compute the R2 score
    r2 = r2_score(y, model.predict(X))
    
    # Return the resulting score
    return r2
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
# Compute the score
score = r2_score(lr, train_X, train_y)

# Test the score
assert_almost_equal(0.010956, score, places=2)
```

-----

## Problem 4: Model Persistence

Complete the function `persist_model`, which accepts two parameters: `name` and `model`. This function will persist the provided `model` into a new file specified by `name`. To persist the machine learning model, you should use the joblib library. By using the function template provided below, you must explicitly:
- Open the file by using the provided name for writing.
- Save the model to this file.

-----

<!-- grade_id: p4_answer, points: 0, grade: false, solution: true -->
```python
import os
import joblib
#from sklearn.externals import joblib
from joblib import dump

def persist_model(name, model):
    '''
    Write a model to the specified file.

    Parameters
    ---------   
    name: A string containg the filename to which the model should be written.
    model: The model that should be saved to the file.

    Returns
    -------
    Nothing
    '''
   
    ### YOUR CODE HERE
    
    # Open the file by using the provided name for writing
    with open(name, 'wb') as file:
        # Save the model to the file
        dump(model, file)
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
# Save model to temporary file
persist_model('test_model.pkl', lr)

# Does the file exist?
assert_true(os.path.exists('test_model.pkl'))

# Test model recreation
with open('test_model.pkl', 'rb') as fin:
    new_model = joblib.load(fin)
    assert_equal(new_model.fit_intercept, True)

# Remove the temporary file
!rm test_model.pkl
```

Output:
```
'rm' is not recognized as an internal or external command,
operable program or batch file.
```

**&copy; 2017: Robert J. Brunner at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode 